<img style="float: left;" src="earth-lab-logo-rgb.png" width="150" height="150" />

# Homework Template: Earth Analytics Python Course: Spring 2020

Before submitting this assignment, be sure to restart the kernel and run all cells. To do this, pull down the Kernel drop down at the top of this notebook. Then select **restart and run all**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below.

* IMPORTANT: Before you submit your notebook, restart the kernel and run all! Your first cell in the notebook should be `[1]` and all cells should run in order! You will lose points if your notebook does not run. 

For all plots and code in general:

* Add appropriate titles to your plot that clearly and concisely describe what the plot shows (e.g. time, location, phenomenon).
* Be sure to use the correct bands for each plot.
* Specify the source of the data for each plot using a plot caption created with `ax.text()`.
* Place ONLY the code needed to create a plot in the plot cells. Place additional processing code ABOVE that cell (in a separate code cell).

Make sure that you:

* **Only include the package imports, code, data, and outputs that are CRUCIAL to your homework assignment.**
* Follow PEP 8 standards. Use the `pep8` tool in Jupyter Notebook to ensure proper formatting (however, note that it does not catch everything!).
* Keep comments concise and strategic. Don't comment every line!
* Organize your code in a way that makes it easy to follow. 
* Write your code so that it can be run on any operating system. This means that:
   1. the data should be downloaded in the notebook to ensure it's reproducible.
   2. all paths should be created dynamically using the os package to ensure that they work across operating systems. 
* Check for spelling errors in your text and code comments


```python
NAME = ""
COLLABORATORS = ""
```

![Colored Bar](colored-bar.png)

# Week 09 Homework - Multispectral Remote Sensing II


## Include the Plots, Text and Outputs Below

For all plots:

* Add appropriate titles to your plot that clearly and concisely describe what the plot shows.
* Be sure to use the correct bands for each plot.
* Specify the source of the data used for each plot in a plot caption using `ax.text()`.


## Project Introduction (10 points)

Read the overview of the cold springs fire: https://www.earthdatascience.org/courses/use-data-open-source-python/data-stories/cold-springs-wildfire/

In the Markdown cell below, add a 2-4 sentence description of the Cold Springs Fire. This should 
include the event:
1. name, 
2. type, 
3. duration / dates and 
4. location. 

<!-- grade_id: cell-a6c72b4cbb716c7c, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

![Colored Bar](colored-bar.png)

```python
# Autograding imports - do not modify this cell
import matplotcheck.notebook as nb
import matplotcheck.autograde as ag
import matplotcheck.raster as rs
```

<!-- grade_id: cell-cade46b91b256c60, points: 5, grade: true, solution: true -->
```python
# Import libraries (5 points) only include imports required to run this notebook
 
# YOUR CODE HERE
raise NotImplementedError()
```


# Define Functions To Process Landsat Data 

For next week (March 18) create the 3 functions below used to process landsat data.
For all functions:

Add a doc string that uses numpy style format. The docstring should include:

    * A one sentence description of what the function does
    * Description of each input variable following numpy docstring standards
    * Description of each output object


## Function 1: crop_stack_data function

Write a function called `crop_stack_data` that: 
1. Takes a **list** of raster TIF files(files_to_crop) and crops all of the files in the list to a given input boundary in GeoPandas GeoDataFrame format.
    * **3 inputs:** 1) list of files, 2) directory to export cropped files, 3) GeoPandas GeoDataFrame to crop the data
2. Returns a stacked **numpy array** of the cropped data and the metadata.
    * **2 outputs:** 1) **numpy array** and 2) **dictionary** containing metadata 

<!-- grade_id: cell-ce1d604c026e2611, points: 0, grade: false, solution: true -->
```python
# Add your function here. do NOT modify the function name
def crop_stack_data(files_to_crop, crop_dir_path, crop_bound):

    # YOUR CODE HERE
    raise NotImplementedError()
```

## Function 2: mask_data

In the call below, write a function called `mask_data` that: 
1. Masks a numpy array using the Landsat Cloud QA layer. 
    * **2 inputs:** numpy array to be masked, 2) landsat qa layer in numpy array format 
2. Returns a masked array. 
    * **1 output:** masked numpy array

<!-- grade_id: cell-2f2fda206fb36d0c, points: 0, grade: false, solution: true -->
```python
# Define mask_data function

def mask_data(arr, path_to_qa):

    # YOUR CODE HERE
    raise NotImplementedError()
```

## Function 3: classify_dnbr

In the cell below, write a function called `classify_dnbr` that: 
1. Masks a numpy array using the Landsat Cloud QA layer. 
    * **1 input:** numpy array containing dNBR data in numpy array format 
2. Returns a classified numpy array. 
    * **1 output:** numpy array with classified values (integers)

<!-- grade_id: cell-033ba180dbc20c63, points: 0, grade: false, solution: true -->
```python
def classify_dnbr(arr):

    # YOUR CODE HERE
    raise NotImplementedError()
```

## Assignment Week 2 - BEGINS March 18 
You will write the function below next week after learning more about MODIS 
h4 data in class on March 18, 2020.

### Function 4: stack_modis_bands

Write a function called `stack_modis_bands` that: 
1. Loops through an `h4` file to collect and stack all "band" layers in the file.
2. Crops each band to the extent of a given input boundary in GeoPandas GeoDataFrame format.
3. Returns a stacked numpy array of the cropped data and the metadata.

#### Function Inputs and Outputs

1. Takes a path to an h4 file and returns the reflectance bands cropped to a specific extent 
    * **2 inputs:** 1) string path to hdf h4 file, 2) crop_bound in GeoDataFrame format
2. Returns a classified numpy array. 
    * **2 outputs:** 1) numpy array containing cropped MODIS reflectance bands, 2) metadata for the numpy array

<!-- grade_id: cell-e730c0516d588d6c, points: 0, grade: false, solution: true -->
```python
def stack_modis_bands(h4_path, crop_bound):

    # YOUR CODE HERE
    raise NotImplementedError()
```

![Colored Bar](colored-bar.png)

# Figure 1 Overview - Grid of 3 Color InfraRed (CIR) Plots: NAIP, Landsat and MODIS

**You will be able to complete the MODIS subplot after March 18 class!**

Create a single figure that contains a grid of 3 plots of color infrared (also called false color) composite images using:

* Post Fire NAIP data (this is the data that you downloaded for your week 6 homework)
* Post Fire Landsat data (use: `et.data.get_data(url="https://ndownloader.figshare.com/files/21941085")`
* Post Fire MODIS data (use: `et.data.get_data('cold-springs-modis-h5')`

For each map, be sure to:

* Crop the data to the fire boundary extent.
* Overlay the fire boundary layer (`vector_layers/fire-boundary-geomac/co_cold_springs_20160711_2200_dd83.shp`).
* Use the band combination **r = infrared band**, **g = red band**, **b = green** band.
* Be sure to label each plot with the data type (NAIP vs. Landsat vs. MODIS) and spatial resolution.

HINT: In a CIR image, the NIR band is plotted on the “red” band, the red band is plotted on the "green" band and the green band is plotted on the "blue" band.

<!-- grade_id: cell-24a209b218bbe2bc, points: 0, grade: false, solution: true -->
```python
# Open fire boundary in this cell

# YOUR CODE HERE
raise NotImplementedError()
```

## Process Landsat Data

In the cells below, open and process your Landsat Data using loops and string manipulation of paths. 

Use the functions `crop_stack_data()` and `mask_data()` that you wrote above to open, 
crop and stack each Landsat data scene (pre and post fire).

#### Notes

If you were implementing this workflow for more scenes, you could write 
a helper function that tested the crs of the crop extent. If it needed to be 
reprojected you could do so and write it out as a file, or store is in a dictionary
for for reuse 
in your loop. For this assignment rather than introducing additional tasks
we will keep it simple and open up and reproject the boundary once.

One way to do this is to create a connection to one single tif file using
rasterio. Once you have the `src` object you can grab the crs and reproject
the fire boundary. The code to do that is below:

```
with rio.open(glob(all_dirs[0] + "/*.tif")[0]) as src:
    fire_bound_utmz13 = fire_boundary.to_crs(src.crs)
```


<!-- grade_id: cell-736f55ecddd6ba9e, points: 0, grade: false, solution: true -->
```python
# Create loop to process Landsat data in this cell
# Use the functions that you wrote above!

# YOUR CODE HERE
raise NotImplementedError()
```

## Landsat Function Tests

<!-- grade_id: cell-27e4ee294fcfa584, points: 5, grade: true, solution: false -->
```python
# DO NOT MODIFY - test: crop_stack_data function
```

<!-- grade_id: cell-a44324eaa47cb6dd, points: 5, grade: true, solution: false -->
```python
# DO NOT MODIFY - test mask_data function

```

<!-- grade_id: cell-f555d110a2c0ba38, points: 5, grade: true, solution: false -->
```python
# DO NOT MODIFY -- Test classify_dnbr function
# BEGIN HIDDEN TEST
### DO NOT REMOVE CODE BELOW ###

points = 0
try:

    dnbr = classify_dnbr(dnbr_landsat_arr)
    assert(isinstance(dnbr, np.ndarray))
    points += 2.5
    print("dnbr_class_bins is a list!")
except AssertionError:
    print("dnbr_class_bins should be a list.")

try:
    assert(set(np.unique(dnbr)) == {1, 2, 3, 4, 5})
    points += 2.5
    print("dnbr_class_bins has the correct values!")
except AssertionError:
    print("dnbr_class_bins does not have the correct values.")

points

# END HIDDEN TEST
```

## Process NAIP Post Fire Data

In the cell below, open and crop the post-fire NAIP data that you downloaded 
for homework 6.

<!-- grade_id: cell-8c391e31f8b2843f, points: 0, grade: false, solution: true -->
```python
# Process NAIP data

# YOUR CODE HERE
raise NotImplementedError()
```

## Process MODIS h4 Data - March 18th, 2020

In the cells below, open and process your MODIS hdf4 data using loops and 
string manipulation of paths. You will learn more about working with MODIS 
in class on March 18th.

Use the function `stack_modis_bands()` that you will write to open and crop 
the MODIS data.


<!-- grade_id: cell-244689f994e59daa, points: 0, grade: false, solution: true -->
```python
# Process MODIS Data

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-41e97b2dd5a5c5d5, points: 10, grade: true, solution: false -->
```python
# DO NOT MODIFY THIS CELL - autograding tests for MODIS function stack_modis_bands

# BEGIN HIDDEN TEST

# Test their modis function works as expected
# We may need to add more specifics about how to create it

modis_path_dirs = glob(os.path.join("cold-springs-modis-h5") + "/*/")

modis_dict = {}
modis_veg_index = {}
points = 0
for adir in modis_path_dirs:

    # Get date for data
    modis_date = os.path.basename(os.path.normpath(adir))
    modis_h4_path = glob(os.path.join(adir, "*.hdf"))[0]
    try:
        modis_arr, modis_meta = stack_modis_bands(modis_h4_path,
                                                  fire_boundary)
        if (type(modis_arr) == np.ndarray):
            points += 5
            print("Excellent, the first object returned from your "
                  "stack_modis_bands function for the dir: {0} "
                  "is a numpy array.".format(adir))
    except:
        print("Oops looks like your stack_modis_bands function doesn't work as"
              "expected. It should return a numpy array")

points
# END HIDDEN TEST
```

## Figure 1: Plot CIR for NAIP, Landsat and MODIS Using Post Fire Data

In the cell below, create a figure with 3 subplots stacked vertically.

In each subplot, plot a CIR composite image using 

NAIP  (first figure axis), 
Landsat (second figure axis), and 
MODIS (third figure axis)

post-fire data on the first, second, and third axes of this figure.

<!-- grade_id: cell-4feb05c88bb2f0d7, points: 0, grade: false, solution: true -->
```python
# Plot CIR of Post Fire NAIP, Landsat & MODIS together in one figure

# YOUR CODE HERE
raise NotImplementedError()
### DO NOT REMOVE LINE BELOW ###
plot01_CIR_res_comparison = nb.convert_axes(plt, which_axes="all")
```

<!-- grade_id: cell-3c3e6bc76ded28b5, points: 15, grade: true, solution: false -->
```python
# DO NOT TOUCH THIS CELL - autograding tests for NAIP subplot

```

<!-- grade_id: cell-0259f2c506f88113, points: 15, grade: true, solution: false -->
```python
# DO NOT TOUCH THIS CELL - autograding tests for Landsat subplot

```

<!-- grade_id: cell-158bfb775a0eddae, points: 15, grade: true, solution: false -->
```python
# DO NOT TOUCH THIS CELL - autograding tests for MODIS subplot

```

![Colored Bar](colored-bar.png)

# Figure 2: Difference NDVI (dNDVI) Using Landsat & MODIS Data

Plot the NDVI difference using before and after Landsat and MODIS data 
that cover the Cold Springs Fire study area. For each axis be sure to:

1. overlay the fire boundary (`vector_layers/fire-boundary-geomac/co_cold_springs_20160711_2200_dd83.shp`) on top of the data.
2. Be sure that the data are cropped using the fire boundary extent.

In the cell below, create a figure with 2 subplots stacked vertically.
Plot dNDVI for Landsat on the first axis and dNDVI using MODIS on the 
second axe of the figure.

Use the before and after data that you processed above to calculate NDVI difference
for both MODIS and Landsat


## NDVI Difference

To create the NDVI Difference raster using before and after fire 
Landsat and MODIS data, you must first calculate NDVI for each 
dataset before and after the fire. 

Once you have the before and after NDVI arrays, you can subtract 
the pre-fire NDVI array FROM the post-fire NDVI array (post-fire minus pre-fire). 

The resulting array will show you change in the area's NDVI from the first image to the second image.

HINT: Remember, you can use `es.normalized_diff(band_1, band_2)` to get the NDVI of an image. 

<!-- grade_id: cell-04c708968d254a97, points: 0, grade: false, solution: true -->
```python
# Plot Difference NDVI for Landsat & MODIS together in one figure

# YOUR CODE HERE
raise NotImplementedError()

### DO NOT REMOVE LINE BELOW ###
plot02_landsat_modis_ndvi_diff = nb.convert_axes(plt, which_axes="all")
```

<!-- grade_id: cell-9b92af5e76ecf4b8, points: 20, grade: true, solution: false -->
```python
# Ignore this cell - autograding tests for Landsat subplot

```

<!-- grade_id: cell-88f2a83e96215ddb, points: 20, grade: true, solution: false -->
```python
# Ignore this cell - autograding tests for MODIS subplot NDVI diff

```

![Colored Bar](colored-bar.png)

# Figure 3 Overview: Difference NBR (dNBR) Using Landsat  & MODIS Data

Create a figure that has two subplots stacked vertically using the same MODIS and Landsat data that you processed above. 

Subplot one: classified dNBR using Landsat data
Subplot two: classified dNBR using MODIS data 

For each subplot, overlay the fire extent boundary `vector_layers/fire-boundary-geomac/co_cold_springs_20160711_2200_dd83.shp`
on top of the dNBR map

To classify each dNBR raster, use the `classify_dnbr()` function that you 
defined above. 

When you plot your MODIS data, you may notice that the data does not contain all of the classes that Landsat contains which can range from 1-5. To ensure that your colormap plots properly, set the `vmin=` and `vmax=` parameters to 1 and 5 respectively when you call `ep.plot_bands()`:

`vmin=1, vmax=5`


## Figure Legend

You only need one legend for this figure. The `ep.draw_legend()` function will create a legend of "boxes" if you provide it with an:

1. `imshow()` image object
2. classes : a list of numbers that represent the classes in your numpy array
3. titles: a list of dNBR class names example: `["High Severity", "Low Severity"]`

### dNBR Classes

Note: if you scaled your data, you may need to scale the values below by a factor of 10.

| SEVERITY LEVEL  | dNBR RANGE |
|----------------|--------------|
| Enhanced Regrowth | < -.1 |
| Unburned        | -.1 to +.1 |
| Low Severity     | +.1 to +.27 |
| Moderate Severity  | +.27 to +.66 |
| High Severity    | > .66 |


HINT: Your dNBR classification list should look like this:
`[-np.inf, -.1, .1, .27, .66, np.inf]`

HINT 2: If you want to use them, these are the colors used in the maps on the website:

`["g", "yellowgreen", "peachpuff", "coral", "maroon"]`

The code to create a custom colormap for the plot is below:

```
nbr_colors = ["g", "yellowgreen", "peachpuff", "coral", "maroon"]
nbr_cmap = ListedColormap(nbr_colors)
```

<!-- grade_id: cell-296ed0f3c3b010fd, points: 0, grade: false, solution: true -->
```python
# Calculate dNBR for Landsat - be sure to use the correct bands!

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1b1b71c7a1d9c6e6, points: 0, grade: false, solution: true -->
```python
# Plot Difference NBR (dNBR) for Landsat & MODIS together in one figure

# YOUR CODE HERE
raise NotImplementedError()

### DO NOT REMOVE LINE BELOW ###
plot03_landsat_dnbr = nb.convert_axes(plt, which_axes="all")
```

<!-- grade_id: cell-39ad421822088e36, points: 25, grade: true, solution: false -->
```python
# Ignore this cell - autograding tests for Landsat subplot

```

<!-- grade_id: cell-bf37c466c5974606, points: 25, grade: true, solution: false -->
```python
# Ignore this cell - autograding tests for MODIS subplot

```

![Colored Bar](colored-bar.png)

# Landsat vs MODIS  Burned Area

In the cell below, print the total area burned in classes 4 and 5 (moderate to high severity) for both datasets 
(Landsat and MODIS).

HINT: Feel free to experiment with loops to complete this part of the homework. 

<!-- grade_id: cell-e8ccdd40f3954765, points: 0, grade: false, solution: true -->
```python
# Total Burned Area in Classes 4 and 5 for Landsat and MODIS

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b8a342e0efca8f12, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

![Colored Bar](colored-bar.png)

# Do not edit this cell!

* Each figure specifies the source of the data (for each plot) using a plot caption created with `ax.text()`.

# Do not edit this cell!

The notebook will also be checked for overall clean code requirements as specified at the **very top** of this notebook! Some of these requirements include (review the top cells for more specifics): 

* Notebook begins at cell [1] and runs on any machine in its entirety.
* PEP 8 format is applied throughout (including lengths of comment and code lines).
* No additional code or imports in the notebook
* Notebook is fully reproducible. This means:
   * reproducible paths using the os module.
   * data downloaded using code in the notebook.
   * all imports at top of notebook.

# DO NOT EDIT

All functions contain docstrings with inputs and outputs clearly indicated

