<img style="float: left;" src="earth-lab-logo-rgb.png" width="150" height="150" />

# Homework Template: Earth Analytics Python Course: Spring 2020

Before submitting this assignment, be sure to restart the kernel and run all cells. To do this, pull down the Kernel drop down at the top of this notebook. Then select **restart and run all**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below.

* IMPORTANT: Before you submit your notebook, restart the kernel and run all! Your first cell in the notebook should be `[1]` and all cells should run in order! You will lose points if your notebook does not run. 

For all plots and code in general:

* Add appropriate titles to your plot that clearly and concisely describe what the plot shows (e.g. time, location, phenomenon).
* Be sure to use the correct bands for each plot.
* Specify the source of the data for each plot using a plot caption created with `ax.text()`.
* Place ONLY the code needed to create a plot in the plot cells. Place additional processing code ABOVE that cell (in a separate code cell).

Make sure that you:

* **Only include the package imports, code, data, and outputs that are CRUCIAL to your homework assignment.**
* Follow PEP 8 standards. Use the `pep8` tool in Jupyter Notebook to ensure proper formatting (however, note that it does not catch everything!).
* Keep comments concise and strategic. Don't comment every line!
* Organize your code in a way that makes it easy to follow. 
* Write your code so that it can be run on any operating system. This means that:
   1. the data should be downloaded in the notebook to ensure it's reproducible.
   2. all paths should be created dynamically using the os package to ensure that they work across operating systems. 
* Check for spelling errors in your text and code comments


```python
NAME = ""
COLLABORATORS = ""
```

![Colored Bar](colored-bar.png)

# Week 09 Homework - Multispectral Remote Sensing II


## Include the Plots, Text and Outputs Below

For all plots:

* In  markdown cell BELOW each plot, add 2-4 sentences that describe what each plot shows. 
* Add appropriate titles to your plot that clearly and concisely describe what the plot shows.
* Be sure to use the correct bands for each plot.
* Specify the source of the data used for each plot in a plot caption using `ax.text()`.

## Project Introduction

In the markdown cell below, adda 2-4sentence description of the Cold Springs Fire. This should include the event:
1. name, 
2. type, 
3. duration / dates and 
4. location. 


<!-- grade_id: cell-a6c72b4cbb716c7c, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

##### dNBR burn classes

If you scaled your data, you may need to scale the
values below by a factor of 10.

| SEVERITY LEVEL  | dNBR RANGE |
|----------------|--------------|
| Enhanced Regrowth | < -.1 |
| Unburned        | -.1 to +.1 |
| Low Severity     | +.1 to +.27 |
| Moderate Severity  | +.27 to +.66 |
| High Severity    | > .66 |

****

![Colored Bar](colored-bar.png)

```python
# Autograding imports - do not modify this cell
import matplotcheck.notebook as nb
import matplotcheck.autograde as ag
import matplotcheck.raster as rs
```

<!-- grade_id: cell-cade46b91b256c60, points: 0, grade: true, solution: true -->
```python
# Import libraries

# YOUR CODE HERE
raise NotImplementedError()


### DO NOT REMOVE LINE BELOW ###
```

![Colored Bar](colored-bar.png)

## Plot 1 - Grid of Plots: NAIP, Landsat and MODIS

Create a single plot that contains a grid of 3, 3 band color infrared (also called false color) plots using:
images using:

* Post Fire NAIP data (this is the data that you already downloaded for your homework earlier this year)
* Post Fire Landsat data (use: `et.data.get_data(url="https://ndownloader.figshare.com/files/21941085")`
* Post Fire MODIS data (use: `et.data.get_data('cold-springs-modis-h5')`

For each map be sure to:

* Crop the data to the fire boundary extent
* Overlay the fire boundary layer (`vector_layers/fire-boundary-geomac/co_cold_springs_20160711_2200_dd83.shp`).
* Use the band combination **r = infrared band**, **g = red band**, **b = green** band.
* Be sure to label each plot with the data type (NAIP vs. Landsat vs. MODIS) and spatial resolution.

Use this figure to help answer question 1 above.
An example of what this plot should look like (without all of the labels that
you need to add), is below.

HINT: In a CIR image, the NIR band is plotted on the “red” band, the red band is plotted using green and the green band is plotted using blue.

<!-- grade_id: cell-033ba180dbc20c63, points: 0, grade: false, solution: true -->
```python
def crop_stack_data(files_to_crop, crop_path, crop_bound):

    # YOUR CODE HERE
    raise NotImplementedError()


def mask_data(arr, path_to_qa):

    # YOUR CODE HERE
    raise NotImplementedError()


def classify_dnbr(arr):

    # YOUR CODE HERE
    raise NotImplementedError()
```

## Process Landsat Data

In the cell or cells below open and process your Landsat Data


# THERE"S NO GOOD WAY TO HANDLE THE FIRE BOUNDARY REPROJECT

option 1: in the loop. here you reprojecting multiple times
option 2: read one tif file and reproject once. it's a manual process.

<!-- grade_id: cell-24a209b218bbe2bc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-736f55ecddd6ba9e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## ALL FUNCTIONS SHOULD BE MOVED TO THE TOP

In the cell below, create a function that does the following: 

1. Loops through an `h4` file and collects and stacks all "band" layers in the file
2. Crops each band to the extent of a given input boundary in geopandas format
3. Returns a stacked numpy array and the metadata for a single band 

Name your function `stack_modis_bands`.

<!-- grade_id: cell-6e2a35a0950b2b06, points: 0, grade: false, solution: true -->
```python

def stack_modis_bands(h4_path, crop_bound):

    # YOUR CODE HERE
    raise NotImplementedError()
```

## Process MODIS Data

In the cell or cells below open and process your MODIS data using loops and string manipulation 
of paths. 

Use the function `stack_modis_bands()` that you wrote above to open and crop the MODIS data for before and after the fire.



<!-- grade_id: cell-244689f994e59daa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-41e97b2dd5a5c5d5, points: 5, grade: true, solution: false -->
```python
# BEGIN HIDDEN TEST

# Test their modis function works as expected
# We may need to add more specifics about how to create it

modis_path_dirs = glob(os.path.join("cold-springs-modis-h5") + "/*/")

modis_dict = {}
modis_veg_index = {}
points = 0
for adir in modis_path_dirs:

    # Get date for data
    modis_date = os.path.basename(os.path.normpath(adir))
    modis_h4_path = glob(os.path.join(adir, "*.hdf"))[0]
    try:
        modis_arr, modis_meta = stack_modis_bands(modis_h4_path,
                                                  fire_boundary)
        if (type(modis_arr) == np.ndarray):
            points += 3
            print("Excellent, the first object returned from your stack_modis_bands function is a numpy array.")
    except:
        print("Oops looks like your function doesn't work as expected")

points
# END HIDDEN TEST
```

## NAIP Data
In the cell below open and crop your NAIP data

<!-- grade_id: cell-8c391e31f8b2843f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

![Colored Bar](colored-bar.png)

## Figure 1: Plot CIR Images for NAIP, Landsat and MODIS Using Post Fire Data

In the cell below, create a figure with 3 subplots stacked vertically.
Plot NAIP, Landsat and MODIS post-fire data on the first, second and third axes of this figure.

<!-- grade_id: cell-4feb05c88bb2f0d7, points: 0, grade: false, solution: true -->
```python
# Plot Post Fire NAIP, Landsat & MODIS together in one Figure

# YOUR CODE HERE
raise NotImplementedError()
### DO NOT REMOVE LINE BELOW ###
plot01_CIR_res_comparison = nb.convert_axes(plt, which_axes="all")
```

<!-- grade_id: cell-3c3e6bc76ded28b5, points: 15, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0259f2c506f88113, points: 15, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-158bfb775a0eddae, points: 15, grade: true, solution: false -->
```python

```

# Figure 2: Landsat and MODIS NDVI Difference 

Create a map of the NDVI difference before and after the Cold Springs Fire. To do this, 
you will need to download cloud free Landsat data from before the Cold Springs fire and crop it 
to the fire extent. Make sure you overlay the fire extent `vector_layers/fire-boundary-geomac/co_cold_springs_20160711_2200_dd83.shp` on top of the NDVI Difference map. 

Be sure to use cloud free data.

Create the same NDVI Difference map using MODIS data before and
after the Cold Springs fire. Overlay the fire extent `vector_layers/fire-boundary-geomac/co_cold_springs_20160711_2200_dd83.shp` on top of the NDVI Difference map. 

### Landsat Data
Use cloud free data that you downloaded for pre fire Landsat calculations. Make sure 
that all the Landsat data is cropped to the fire boundary extent before you perform
any calculations on it.

### MODIS Data

Crop the MODIS data to the fire boundary extent.

### NDVI Differencing

To get a map of the NDVI Difference from before and after the fire, you must first create NDVI 
rasters of the study area before and after the fire. Once you have those arrays, to get the 
Difference subtract the pre-fire array from the post-fire array. The resulting array will 
show you change in the area's NDVI from the first image to the second image. The operation
itself looks like this: `ndvi_difference = post_fire_ndvi - pre_fire_ndvi`.
****

HINT: Remember, you can use `es.normalized_diff(band1, band2)` to get the NDVI of an image. 

<!-- grade_id: cell-04c708968d254a97, points: 0, grade: false, solution: true -->
```python
# Difference NDVI before and after the fire Landsat and MODIS


# YOUR CODE HERE
raise NotImplementedError()

### DO NOT REMOVE LINE BELOW ###
plot02_landsat_modis_ndvi_diff = nb.convert_axes(plt, which_axes="all")
```

<!-- grade_id: cell-9b92af5e76ecf4b8, points: 25, grade: true, solution: false -->
```python
# add tests here

```

<!-- grade_id: cell-88f2a83e96215ddb, points: 25, grade: true, solution: false -->
```python

```

![Colored Bar](colored-bar.png)

## Figure 3: Difference NBR (dNBR) Using Landsat  & MODIS Data

Create a map of the classified dNBR using Landsat data collected before and
after the Cold Springs fire. Overlay the fire extent `vector_layers/fire-boundary-geomac/co_cold_springs_20160711_2200_dd83.shp`
on top of the dNBR map. 

Create the same classified dNBR map using MODIS data before and
after the Cold Springs fire. Overlay the fire extent `vector_layers/fire-boundary-geomac/co_cold_springs_20160711_2200_dd83.shp`
on top of the dNBR map. 


Add one legend for both plots that clearly describes each burn "class"
that you applied using the table provided in the lessons. The classes will be the same values as it is an index.


#### dNBR burn classes

The burn classes that you should use to classify the data are below:


Note that depending on how you scaled your data, you may need to scale the
values below by a factor of 10.

| SEVERITY LEVEL  | dNBR RANGE |
|-----------------|-------------|
| Enhanced Regrowth | < -.1 |
| Unburned       | -.1 to +.1 |
| Low Severity     | +.1 to +.27 |
| Moderate Severity  | +.27 to +.66 |
| High Severity   | > .66 |

****


HINT: Your dNBR classification list should look like this:
`[-np.inf, -.1, .1, .27, .66, np.inf]`

HINT 2: If you want to use them, these are the colors used in the maps on the website:

`["g", "yellowgreen", "peachpuff", "coral", "maroon"]`

The code to create a custom colormap for the plot is below:

```
nbr_colors = ["g", "yellowgreen", "peachpuff", "coral", "maroon"]
nbr_cmap = ListedColormap(nbr_colors)
```

## Create NBR Bins Object 

In the cell below create the object that you will use to define bins 
to classify both Landsat and MODIS data. Name your object:

`dnbr_class_bins`

<!-- grade_id: cell-bd60016448bbe655, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f555d110a2c0ba38, points: 5, grade: true, solution: false -->
```python
# Here add a tests to ensure
# 1. dnbr class bins is a list and
# 2. is has the correct values - i suggest just testing that the 2nd value is
# -1 and the 4th is .66
# the tests can be worth 5 points total

# Here are some sanity checks that you set up your dnbr_class_bins correctly
# The first test ensures that the data type is correct, i.e., that dnbr_class_bins
# is a list. The second test checks two of the values in the list to make sure that
# the correct values are in the correct order in the list.

# These tests ARE graded, so make sure you're passing both of the tests before you move on!

### DO NOT REMOVE CODE BELOW ###

points = 0
try:
    assert(isinstance(dnbr_class_bins, list))
    points += 2
    print("dnbr_class_bins is a list!")
except AssertionError:
    print("dnbr_class_bins should be a list.")

try:
    assert(dnbr_class_bins[1] == -.1 and dnbr_class_bins[4] == .66)
    points += 3
    print("dnbr_class_bins has the correct values!")
except AssertionError:
    print("dnbr_class_bins does not have the correct values.")

points
```

<!-- grade_id: cell-296ed0f3c3b010fd, points: 0, grade: false, solution: true -->
```python
# Calculate dNBR for Landsat - be sure to use the correct bands!

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9cf18a3835896598, points: 0, grade: false, solution: true -->
```python
# Calculate MODIS dNBR

# YOUR CODE HERE
raise NotImplementedError()
```

![Colored Bar](colored-bar.png)

## Homework Figure 3: Plots of Landsat & MODIS Difference Normalized Burn Ratio  dNBR

When you plot your MODIS data, you may notice that the data does not contain all of the classes that Landsat contains which could range from 1-5. To ensure that your colormap plots properly, set the vmin and vmax parameters to 1 and 5 respectively when you call `ep.plot_bands()`.

`vmin=1, vmax=5`

### Plot Legend

You only need one legend for this figure. The es draw_legend function will create a legend of "boxes" if you provide it with an 

1. `imshow()` image object
2. classes : a list of numbers that represent the classes in your numpy array
3. titles: a list of class names example: `["High Severity", "Low Severity"]`

<!-- grade_id: cell-1b1b71c7a1d9c6e6, points: 0, grade: false, solution: true -->
```python
# Plot dNBR using Landsat

# YOUR CODE HERE
raise NotImplementedError()
### DO NOT REMOVE LINE BELOW ###
plot03_landsat_dnbr = nb.convert_axes(plt, which_axes="all")
```

<!-- grade_id: cell-39ad421822088e36, points: 25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-bf37c466c5974606, points: 25, grade: true, solution: false -->
```python

```

![Colored Bar](colored-bar.png)

# Landsat vs MODIS  Burned Area

Below, print the total area burned in classes 4 and 5 (moderate to high severity) for both datasets.

HINT: Feel free to experiment with loops to complete this part of the homework. 


<!-- grade_id: cell-e8ccdd40f3954765, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```